```python
import hashlib # for grading

# Standard imports
import numpy as np
from numpy import inf
import pandas as pd
from collections import Counter, OrderedDict
import re
import string
import math
from collections import Counter
import warnings; warnings.simplefilter('ignore')

# NLTK imports
from nltk.tokenize import WordPunctTokenizer
from nltk.stem.snowball import SnowballStemmer
from nltk.corpus import stopwords

# SKLearn related imports
import sklearn
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.feature_extraction.text import TfidfTransformer
from sklearn.pipeline import Pipeline
from sklearn.base import TransformerMixin
from sklearn import preprocessing

from sklearn.naive_bayes import MultinomialNB
from sklearn.model_selection import train_test_split

# Load the dataset
df = pd.read_csv('data/imdb_sentiment.csv')

# Get the text
docs = df['text']

# Split in train and validation
train_df, validation_df = train_test_split(df, test_size=0.2, random_state=42)
```

## Q1.

You were a given the task of answering some questions about a file with the 10000 most commons password from the latest leak your rival corporation has suffered. You don't want to suffer the same fate, so you are trying to come up with a set of restrictions that should reduce the chances of this happening on your side!

<!-- grade_id: cell-a88528290dd35f7d, points: 0, grade: false, solution: true -->
```python
def find_all_in_file(pattern, path):
    """
    Function that returns all matches of a certain pattern in a certain text.
    Args:
    pattern - regex pattern
    path - path to the file
    """
    # YOUR CODE HERE
    return re.findall(pattern,
                      open(path, 'r').read(),
                      re.MULTILINE)
```

#### Q1.a)

Everybody knows that long password tend to be more secure.. Your boss wonders how many of those common password were even 10 or more characters. Get him the number of passwords that have 10 or more characters.

<!-- grade_id: cell-7cddd8e2e48afb31, points: 0, grade: false, solution: true -->
```python
# ret = find_all_in_file(...)

# YOUR CODE HERE
ret = find_all_in_file("^[^\s]{10,}$", 
                       "data/10k-most-common.txt")
```

<!-- grade_id: cell-5a358ed83c473214, points: 1, grade: true, solution: false -->
```python
print("Number of password with more than 10 or more characters: ", len(ret))
assert len(ret) == 51
```

Output:
```
Number of password with more than 10 or more characters:  51
```

#### Q1.b)

As expected people are lazy and have short passwords. What about the number of passwords that start with digits? And the number of passwords that don't start with digits? 

(Get the number of passwords that start with digits and without digits.)

<!-- grade_id: cell-447eb675481e47fd, points: 0, grade: false, solution: true -->
```python
# ret = find_all_in_file(...)

# YOUR CODE HERE
ret = find_all_in_file("^[0-9].+$", "data/10k-most-common.txt")
```

<!-- grade_id: cell-535072609d84d577, points: 1, grade: true, solution: false -->
```python
print("Number of passwords that start with digits: " , len(ret))
assert len(ret) == 659
```

Output:
```
Number of passwords that start with digits:  659
```

<!-- grade_id: cell-66af8912995d91a2, points: 0, grade: false, solution: true -->
```python
# ret = find_all_in_file(...)

# YOUR CODE HERE
ret = find_all_in_file("^[^0-9].+$", "data/10k-most-common.txt")
```

<!-- grade_id: cell-1c17b943b63b52ae, points: 1, grade: true, solution: false -->
```python
print("Number of passwords that start without digits: " , len(ret))
assert len(ret) == 9341
```

Output:
```
Number of passwords that start without digits:  9341
```

#### Q1.c)

So, people prefer to start their passwords with other things than digits. That was to be expected too. Your boss is also very curious about the number of passwords that are variations of the word _pass_. Can you get him that number?

(Get the number of passwords that are variants of the word pass. For instance _pass_, _password_, _passpass_, should be matched.)

<!-- grade_id: cell-2bf14a2df3dd3e4e, points: 0, grade: false, solution: true -->
```python
# ret = find_all_in_file(...)

# YOUR CODE HERE
ret = find_all_in_file("^pass[A-Za-z]*$", "data/10k-most-common.txt")

```

<!-- grade_id: cell-19b61469674be299, points: 1, grade: true, solution: false -->
```python
print("Number of variants of \"pass\": " , len(ret))
assert len(ret) == 13
```

Output:
```
Number of variants of "pass":  13
```

#### Q1.d)

Your boss now comes up with the idea that a password should start with 3 numbers and be followed by letters, but only from b to m. You argue that this would make the task of guessing the password even easier. But he doesn't care. The only way to get him out of this idea is to show him that some passwords in this list follows that pattern. Can you find those passwords?

(Get the number of password that start with 3 numbers and that end with any number of letters from b to m.)

<!-- grade_id: cell-e92132a5b63400f3, points: 0, grade: false, solution: true -->
```python
# ret = find_all_in_file(...)

# YOUR CODE HERE
ret = find_all_in_file("^[0-9]{3}[b-m]+$", "data/10k-most-common.txt")

```

<!-- grade_id: cell-5150c254b592778e, points: 1, grade: true, solution: false -->
```python
print("Number of password: " , len(ret))
assert len(ret) == 2
```

Output:
```
Number of password:  2
```

---

## Q2.

For this question we will be looking at the first few pages of a [great book](https://en.wikipedia.org/wiki/1Q84). You will be ask to perform some common pre-processing operations in natural language processing so you can more easily manipulate the data to asnswer the next questions.

#### Q2.a)

The first thing you will have to do is to tokenize the data. We want to get the words of the provided sample of "1Q84", so, choose the tokenizer accordingly. (Import it from nltk).

<!-- grade_id: cell-a8b7930a61806e32, points: 0, grade: false, solution: true -->
```python
def tokenize_all_file(tokenizer, file):
    """
    Returns the a list with the tokens of given text
    
    Args:
    tokenizer - nltk tokenizer
    file - path to the file to tokenize
    """
    
    # Initialize the list
    # text = ...
    # Tokenize each line of the file
    # ...
    # Flatten the list
    # text_flat = ...
    # return text_flat
    
    # YOUR CODE HERE
    text = list()
    
    for line in open(file, "r"):
        text.append(tokenizer.tokenize(line))
    
    text_flat = [w for s in text for w in s]

    return text_flat

```

<!-- grade_id: cell-864f368a29fd9c6b, points: 0, grade: false, solution: true -->
```python
# tokenizer = ...
# path = "../data/murakami.txt"
# text = tokenize_all_file(...)

# YOUR CODE HERE
tokenizer = WordPunctTokenizer()
path = "data/murakami.txt"
text = tokenize_all_file(tokenizer, path)
```

<!-- grade_id: cell-f829c5a222c54690, points: 1, grade: true, solution: false -->
```python
assert len(text) == 2582
```

#### Q2.b)

Now that the data is tokenized you want to answer this question: "How many different non-stopwords words did Murakami use in the beginning pages?" 

NOTE: Stopwords adapted from [here](Stopwords adapted from [here](https://gist.github.com/sebleier/554280). (Notice what we added some specific things, like ?" and ." to the stopwords. This was shown to be a limitation of the nltk tokenizer so it will be removed that way, instead of the more conventional way. This goes to show that there are more powerful tokenizers that you should use in the case you have to perform tokenization in the future.)

<!-- grade_id: cell-8805d7a045441a7c, points: 0, grade: false, solution: true -->
```python
def preprocess_data(tokens, sw_path, filter_sw=False):
    """
    Returns list with the preprocessed data
    Args:
    tokens - tokenized list
    sw_path - path to the subwords
    filter_sw - set to true if you want to remove subwords
    """
    
    #if filter_sw:
        # Create the list of stopwords from the file english_stopwords.txt
        # stopwords = ...
        # Lowercase all the words and make 
        # lc_text = ...
        # Filter the stopwords from the text
        # rem_sw_text = ...
        # Remove the punctuation (take advantage of string.punctuation)
        # filtered_text = ...
    
    #else:
        # Lowercase all the words and make 
        # lc_text = ...
        # Remove the punctuation (take advantage of string.punctuation)
        # filtered_text = ...
    
    #return filtered_text
    
    # YOUR CODE HERE
    if filter_sw:
        # Create the list of stopwords from the file english_stopwords.txt
        stopwords = [line.strip("\n") for line in open(sw_path, "r")]
        # Lowercase all the words and make 
        lc_text = [w.lower() for w in tokens]
        # Filter the stopwords from the text
        rem_sw_text = list(filter(lambda x: x not in stopwords, lc_text))
        # Remove the punctuation (take advantage of string.punctuation)
        filtered_text = list(filter(lambda x: x not in string.punctuation, rem_sw_text))
    
    else:
        # Lowercase all the words and make 
        lc_text = [w.lower() for w in tokens]
        # Remove the punctuation (take advantage of string.punctuation)
        filtered_text = list(filter(lambda x: x not in string.punctuation, lc_text))
    
    return filtered_text

```

<!-- grade_id: cell-5c7c5a8d817108b7, points: 0, grade: false, solution: true -->
```python
# Get the filtered tokens list
# filtered_text = preprocess_data(...)
# Number of different words
# nr_diff_words = ,,,

# YOUR CODE HERE
filtered_text = preprocess_data(text, "data/english_stopwords.txt", filter_sw=True)    
# Number of different words
nr_diff_words = len(set(filtered_text))
```

<!-- grade_id: cell-f94749703218e9af, points: 1, grade: true, solution: false -->
```python
print("Murakami used {} different words (stopwords not included).".format(nr_diff_words))
assert len(filtered_text) == 1119
assert nr_diff_words == 687
```

Output:
```
Murakami used 687 different words (stopwords not included).
```

---

## Q3.

Your goal is now to find the top three most common word stems. You should apply the snowball stemmer at the `filtered_text` you obtained in Q2. Take advantage of the function `Counter` from [collections](https://docs.python.org/2/library/collections.html) to find the six most common stems.

<!-- grade_id: cell-98884174107979f5, points: 0, grade: false, solution: true -->
```python
def stem_list(tokens):
    """
    Returns a list with the stems of a list of tokens
    Args:
    tokens - list with the tokens
    """
    #stemmer = SnowballStemmer(...)
    #stems = ...
    #return stems

    # YOUR CODE HERE
    stemmer = SnowballStemmer("english", ignore_stopwords=True)
    stems = list(map(stemmer.stem, tokens))
    return stems
```

<!-- grade_id: cell-43e8c6c7722e2b0b, points: 0, grade: false, solution: true -->
```python
#stems = stem_list(...)
#Call counter to get the most common 6 stems

# YOUR CODE HERE
stems = stem_list(filtered_text)
Counter(stems).most_common(6)

```

Output:
```
[('aomam', 28),
 ('driver', 18),
 ('like', 13),
 ('name', 13),
 ('could', 11),
 ('would', 11)]
```

<!-- grade_id: cell-7cb834f08de238d9, points: 0, grade: false, solution: true -->
```python
# answer = ... (copy the output of the previous cell. It should be a list with 6 tuples)

# YOUR CODE HERE
answer = [('aomam', 28), ('driver', 18), ('name', 13), ('like', 13), ('could', 11), ('would', 11)]

```

<!-- grade_id: cell-95a3f0176d152730, points: 1, grade: true, solution: false -->
```python
assert answer == [('aomam', 28), ('driver', 18), ('name', 13),
                  ('like', 13), ('could', 11), ('would', 11)]
```

---

## Q4.

#### Q4.a)

Start by implementing your own function to get the list of n-grams from a list of tokens.

<!-- grade_id: cell-43601c5f562f5867, points: 0, grade: false, solution: true -->
```python
def ngrams(n, text):
    """
    Returns list of tuples for all the n-grams
    
    Args:
    n - the n in n-grams
    text: list of tokenized data
    """
    # YOUR CODE HERE
    return [tuple(text[i:i+n]) for i in range(len(text)-n+1)]
```

<!-- grade_id: cell-f601c9c7d4b35092, points: 2, grade: true, solution: false -->
```python
assert ngrams(2, "The driver made a mistake".split()) == [('The', 'driver'), ('driver', 'made'), ('made', 'a'), ('a', 'mistake')]
assert ngrams(3, "The driver made a mistake".split()) == [('The', 'driver', 'made'), ('driver', 'made', 'a'), ('made', 'a', 'mistake')]
assert ngrams(4, "The driver made a mistake".split()) == [('The', 'driver', 'made', 'a'), ('driver', 'made', 'a', 'mistake')]
```

#### Q4.b)

Now return to Murakami's text and, using preprocessed data without removing the stopwords, find the most common bigram, trigram and 4-gram. Once again, take advantage of `Counter` and `most_common()`.

<!-- grade_id: cell-539dba0b0aabb366, points: 0, grade: false, solution: true -->
```python
# filtered_text = preprocess_data(...)
# bigram_ = ...
# trigram_ = ...
# fourgram_ = ...

# YOUR CODE HERE
filtered_text = preprocess_data(text, "")
bigram_ = Counter(ngrams(2, filtered_text)).most_common(1)[0]
trigram_ = Counter(ngrams(3, filtered_text)).most_common(1)[0]
fourgram_ = Counter(ngrams(4, filtered_text)).most_common(1)[0]

```

<!-- grade_id: cell-91094315aa4e5abf, points: 1, grade: true, solution: false -->
```python
assert bigram_ == (('the', 'driver'), 15)
assert trigram_ == ((',"', 'the', 'driver'), 6)
assert fourgram_ == ((',"', 'the', 'driver', 'said'), 4)
```

What is the highest value of n, before the most common n-gram starts to be unique?

<!-- grade_id: cell-566e26fa1479e9ff, points: 0, grade: false, solution: true -->
```python
# n = ...

# YOUR CODE HERE
for n in range(1, 50):
    ans = Counter(ngrams(n, filtered_text)).most_common(1)[0][1]
    if ans > 1: continue
    else: n-=1; break

```

<!-- grade_id: cell-a060844c68988375, points: 1, grade: true, solution: false -->
```python

answer = hashlib.sha256(bytes(n)).hexdigest()

assert answer == "837885c8f8091aeaeb9ec3c3f85a6ff470a415e610b8ba3e49f9b33c9cf9d619"
```

## Q5.

Which of these statements is true?

**A** - Text is structured and so can be immediately used by a classifier, without much preprocessing needed.

**B** - A Bag of Words representation carries syntactic information (i.e. word order matters in the BoW representation).

**C** - Removal of stopwords may or may not be useful to a certain usecase, depending on our choice of features and dataset used.

**D** - The less a word appears in a document and the more it appears in other documents, the higher is the TF-IDF of that word in that document.

<!-- grade_id: cell-24849fb80014cefb, points: 0, grade: false, solution: true -->
```python
#answer = # Answer with a one character string, capitalized


# YOUR CODE HERE
answer = 'C'
```

<!-- grade_id: cell-2ce91f1c34af8032, points: 2, grade: true, solution: false -->
```python
expected_hash = '6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d'
assert hashlib.sha256(str(answer).encode()).hexdigest() == expected_hash
```

## Q6.

Let's get once again a small sample of our dataset and manually vectorize and transform it into a BoW representation.

```python
small_docs = docs[:200]
stopwords = [line.strip("\n") for line in open("data/english_stopwords.txt", "r")]

def build_vocabulary():
    vocabulary = Counter()

    for doc in small_docs:
        words = [word for word in doc.split() if word not in stopwords]
        vocabulary.update(words)
    
    return OrderedDict(vocabulary.most_common())

def vectorize():
    vocabulary = build_vocabulary()
    vectors = []
    for doc in small_docs:
        words = doc.split()
        vector = np.array([doc.count(word) for word in vocabulary if word not in stopwords])
        vectors.append(vector)
    
    return vectors

def build_df():
    return pd.DataFrame(vectorize(), columns=build_vocabulary())

BoW = build_df()
```

Your boss loves TF-IDF so he asks you to implement it. However, he prefers a slightly different formulation from the one we learned before:

$$ tfidf _{t, d} =(tf_{t,d})*(log_2{(1 + \frac{N}{df_{t}})})  $$

Implement his preferred version of TF-IDF.

<!-- grade_id: cell-eb6996cb41762895, points: 0, grade: false, solution: true -->
```python
def new_tfidf(BoW_df):
    """
    Returns pandas dataframe of a tfidf representation from a BoW representation dataframe.

    Args:
    BoW_df - dataframe with document word counts (Bag of Words)
    """
    # tf = (...)
    
    # def _idf(column):
    #   return (...)
    
    # tf_idf = (...)
    
    # return tf_idf

    # YOUR CODE HERE
    tf = BoW.div(BoW.sum(axis=1), axis=0)

    def _idf(column):
        return np.log2(1 + len(column) / sum(column > 0))

    tf_idf = (tf).multiply(tf.apply(_idf))
    
    return tf_idf

```

```python
new_scores = new_tfidf(BoW)
```

<!-- grade_id: cell-5b8f988a69678772, points: 2, grade: true, solution: false -->
```python
assert(math.isclose(new_scores['movie'][0], 0.009717385023827248),
       math.isclose(new_scores['film'][10], 0.019778475747522496),
       math.isclose(new_scores['nice'][16], 0.010851136310680626),
       math.isclose(new_scores['good'][128], 0.00989061193998239))
```

## Q7.

Let's get our `TextCleanerTransformer` and clean our data.

```python
from TextCleanerTransformer import TextCleanerTransformer
# Initialize a tokenizer and a stemmer
tokenizer = WordPunctTokenizer()
stemmer = SnowballStemmer("english", ignore_stopwords=True)
regex_list = [("<[^>]*>", "")
             ]
cleaner = TextCleanerTransformer(tokenizer, stemmer, regex_list)
docs = cleaner.transform(df.text.values)
```

Here's one of the movie reviews in the imdb dataset:

```python
df['text'].values[2013]
```

For this exercise, transform your documents into a matrix of tf-idf scores using sklearn. Then, get the corresponding string of the most important word of this document (with index `2013`) according to TF-IDF.

```python
vectorizer = CountVectorizer()
vectorizer.fit(docs)
word_count_matrix = vectorizer.transform(docs)
tfidf = TfidfTransformer()
tfidf.fit(word_count_matrix)
word_term_frequency_matrix = tfidf.transform(word_count_matrix)
```

<!-- grade_id: cell-bcc7341a8f2e3160, points: 0, grade: false, solution: true -->
```python
def most_relevant_word(word_term_frequency_matrix, idx):
    """
    Returns the word with highest tfidf score in a specified document.
    
    Args:
    word_term_frequency_matrix - tfidf sparse matrix of the dataset
    idx - index of the relevant document
    """
    # max_word_idx = (...)
    # vocab = (...)
    
    # most_relevant_word = (...)
    
    # return most_relevant_word

    # YOUR CODE HERE
    max_word_idx = word_term_frequency_matrix[idx].argmax()
    vocab = vectorizer.vocabulary_
    inv_vocab = {v: k for k, v in vocab.items()}
    most_relevant_word = inv_vocab[max_word_idx]
    
    return most_relevant_word

```

<!-- grade_id: cell-962273848ec48118, points: 2, grade: true, solution: false -->
```python
assert(most_relevant_word(word_term_frequency_matrix, 2013) == 'morena')
```

## Q8.

In this exercise, build a Pipeline to classify a review as positive or negative. Use `MultinomialNB` as your final classifier, train it and get an accuracy score above 86% on the imdb validation dataset, by choosing the best set of parameters of `TextCleanerTransformer()`, `CountVectorizer()` and `TfidfTransformer()`, according to what we learned in Part III.

Hint: Try to use more than unigrams! Also, remember what we said about stopwords and feature space size in Part III of the Learning Notebooks?

```python
# Encode the labels
le = preprocessing.LabelEncoder()
le.fit(train_df['sentiment'].values)

train_df['sentiment'] = le.transform(train_df['sentiment'].values)
validation_df['sentiment'] = le.transform(validation_df['sentiment'].values)
```

<!-- grade_id: cell-e0712ceb628e6ed6, points: 0, grade: false, solution: true -->
```python
def train_and_validate():
    """
    Train a model using sklearn's Pipeline and return it along with its 
    current accuracy in the validation set.
    """
    
    # Build the pipeline
    # text_clf = Pipeline(...)
    
    # Train the classifier
    # (...)

    # predicted = (...)
    # acc = (...)
    
    # return text_clf, acc
    
    # YOUR CODE HERE
    text_clf = Pipeline([('stemm', TextCleanerTransformer(tokenizer, stemmer, regex_list)),
                       ('vect', CountVectorizer(ngram_range=(1,2), max_features=30000)),
                       ('tfidf', TfidfTransformer()),
                       ('clf', MultinomialNB())])

    # Train the classifier
    text_clf.fit(map(str, train_df['text'].values), train_df['sentiment'].values)

    predicted = text_clf.predict(map(str, validation_df['text'].values))
    acc = np.mean(predicted == validation_df['sentiment'])

    return text_clf, acc

```

<!-- grade_id: cell-07b6f3694941ac81, points: 2, grade: true, solution: false -->
```python
_, acc = train_and_validate()
print("Accuracy: {}".format(acc))
assert(acc >= 0.86)
```

Output:
```
Accuracy: 0.869
```

```python

```